---

<h2>Welcome to the first tutorial of Business Lab & Professional Skills!</h2>

---

**Requirements:**

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".
- Fill in your name and student id below (inside of the quotation marks):

```python
NAME = ""
STUDENT_ID = ""
```

---
## Exercise 1: Strings, Lists and Loops

<b>Exercise 1A:</b> Print 'Hello, world!'

<!-- grade_id: helloworld, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

---
<b>Exercise 1B:</b> In Python, anything inside quotes is a string. You can use either single quotes `'…'` or double quotes `"…"` for your strings.

Task: Print the following four sentences exactly as written. Use a separate `print()` statement for each sentence.

Sentences:
- A key strength of Python lies in its readability and simplicity, as famously stated by Tim Peters: "Readability counts."
- Python's versatility and extensive library support make it a robust choice for a wide range of applications, from web development to data analysis.
- Guido van Rossum, the creator of Python stated: "Python is an experiment in how much freedom programmers need."
- When I asked a friend which foreign language she prefers, she said: "Python is my favourite language".

Requirements:
- Use exactly four `print()` calls. One per sentence.
- Preserve punctuation and capitalization exactly as shown.
- If a sentence contains double quotes ("…"), wrap the string in single quotes '…' OR escape the inner double quotes with `\"`.
- If a sentence contains single quotes (`…`), wrap the string in double quotes '…' OR escape the inner single quotes with `\'`.
- Do not combine the sentences into one print statement.



<!-- grade_id: cell-0c14f1f1345bd22c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

---
<b>Exercise 1C:</b> Create a Python list named `friends` that contains the names of four of your friends.

<!-- grade_id: cell-c45d3d6bce1ddf88, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

---
<b>Exercise 1D:</b> Using the list of friend names you created earlier, print each name on its own line using a `for` loop.

```python
# YOUR CODE HERE
```

---
<b>Exercise 1E:</b> Using your existing list of friend names, print a personalized message to each friend saying they too should learn programming with Python. Use an *f-string*.

Guidance:
- Put the letter `f` immediately before the opening quote: `f"..."` or `f'...'`.
- Insert variables or expressions inside curly braces: `{variable}`.

```python
# YOUR CODE HERE
```

---
<b>Exercise 1F:</b> Using the personalized message from Exercise 1E, print a message to each friend where the friend’s name is in UPPERCASE. Keep using an f-string. Use the `.upper()` method on the name inside the f-string, like `{friend.upper()}`.

```python
# YOUR CODE HERE
```

---
<b>Exercise 1G:</b>
Add the names of two more friends to your list of friends. Use either `.append()` to add the names one by one. Or extend your `friends` list with another list using `.extend()`.

```python
# YOUR CODE HERE
```

---
<b>Exercise 1H:</b>
Check how many names are in your list using the built-in `len()` function. Print the length of the list. If the number is not what you expect, revise the code in the cells above and re-run them in order.

```python
# YOUR CODE HERE
```

---
## Exercise 2: Mean Lists

Python contains many built-in functions which are always available without any imports. Two of the most useful for working with collections are `len()` and `sum()`. However, there may be times where you need to construct your own functions. We will explore custom functions in this exercise.

---
<b>Exercise 2a:</b> Create a list called `my_numbers` that contains the following numbers: 5, 20, 8, 12, 9 and 17.

```python
# YOUR CODE HERE
```

**Exercise 2b**: Print the length of the list `my_numbers` using the `len()` function.

```python
# YOUR CODE HERE
```

**Exercise 2c**: Compute and print the sum of all of the numbers in the list `my_numbers` using the `sum()` function.

```python
# YOUR CODE HERE
```

**Exercise 2d**: Compute and print the arithmetic mean of all of the numbers in the list `numbers` by combining the `len()` and the `sum()` functions.

```python
# YOUR CODE HERE
```

<b>Exercise 2e:</b> To make your code easier to use and maintain, you can create your own functions. A function is a reusable block of code that performs a specific task. You define it once, and then you can call it (use it) as many times as you need.

To create your own function, you can start with the following template:


```
def my_function(parameter1, parameter2):
    """
    Optional docstring that explains
    what the code does.
    """

    #Do something with your parameters, e.g. compute the exponent:
    result = parameter1 ** parameter2

    #Return the result:
    return result

```

Note the indentation. This marks what belongs inside the function.


You can now call the function as follows:

```
result=my_function(2,5)
print(result)
```


**Task:** 

1) Define your own function called `mean` that:
    - Takes one parameter (a list of numbers).
    - Calculates the arithmetic mean (average) of those numbers.
    - Returns the result.
2) Call your `mean` function using your list of numbers named `my_numbers`.
    - Store the result in a variable (e.g. `average`)
    - Print the result to check if your function works.


```python
# YOUR CODE HERE
```

<b>Exercise 2f:</b> See what happens if you call the function on an empty list. You can create an empty list as follows:

```
empty_list= []
```

```python
# YOUR CODE HERE
```

<b>Exercise 2g:</b> As you noticed, the function has an issue: if it receives an empty list, it will cause an error.

Create a modified version of your function that checks whether the list is empty before performing any operations. If the list is empty, the function should print an error message (e.g. `'The list is empty'`).

Recall the structure for conditional statements in Python:

```
if condition1:
    # Code to run if condition1 is True
    ...
elif condition2:
    # Code to run if condition2 is True
    ...
else:
    # Code to run if none of the above conditions are True
    ...
```

```python
# YOUR CODE HERE
```

---
## Exercise 3: Conditionals and Functions
<b>Exercise 3A:</b> Create a function called <i><b>is_even</i></b> if a number is even. Return `True` if even, `False` if odd. Remember that you can use the modulus operator (`%`) to check if a number is even or odd.

<!-- grade_id: cell-7c37a40fc60f01f4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b1cbb85cfa0a3fe7, points: 1, grade: true, solution: false -->
```python
#Do not make changes to this cell!
assert is_even(2)==True
assert is_even(11)==False
```

---
<b>Exercise 3B:</b> In mathematics a <b>factorial</b> is the product of all positive integers less than or equal to a given positive integer.</br>
</br>
**Notation:** The factorial of an integer *n* is written as `n!`</br>
</br>
Examples:</br>
    1! = 1</br>
    2! = 1 * 2 = 2</br>
    3! = 1 * 2 * 3 = 6</br>
    4! = 1 * 2 * 3 * 4 = 24</br>
</br>
</br>
**Task:** Create a function named `fact` that computes and returns the factorial for any positive integer.

**Hint:** Iterate over the numbers from 1 to n, storing intermediate results to calculate the final product.

<!-- grade_id: cell-5026eeafc9ff6fee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-2dec7c80a13db806, points: 1, grade: true, solution: false -->
```python
#Do not make changes to this cell!
assert fact(1)==1
assert fact(2)==2
```

---
<b>Exercise 3C:</b> Modify your factorial function to return an error message if the user does not enter a positive integer. The error message should state: 'Please enter a positive integer.'

<!-- grade_id: cell-925a0fc10176d3a8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-274c3cf57d2f6183, points: 1, grade: true, solution: false -->
```python
#Do not make changes to this cell!
```

---
<b>Exercise 3D:</b> Create a function named <i>pyramid</i> that prints a 'half pyramid' for a given number of layers:</br>

For example:</br>
`pyramid(5)` should generate the following 'half pyramid':

```
    #
   ##
  ###
 ####
#####
```

**Hint:** Use a loop to print each row. For row `i` (starting from `1`), print a certain number of spaces followed by a certain number of `#` symbols so that the `#`s are right-aligned. You can use string multiplication to generate both the spaces and the `#`s.

<!-- grade_id: cell-60e59e943af2ae2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pyramid(8)
```

---
## Exercise 4

Write a function that abbreviates someone's first named to a single letter followed by a dot. It should then returns this letter + the person's last name. The function should be called <i>abbreviate</i>.

For example:</br>
Adam Smith --> A. Smith</br>
Milton Friedman --> M. Friedman</br>
Tjalling Koopmans --> T. Koopmans</br>
Christine Lagarde --> C. Lagarde</br>
</br>
Hints:<br>
- You can split a string a into a list of substrings using the `.split()` method.<br>
- You can use a slice `[ ]` to select an element from a list.
- You can convert a string to title case using the `.title()` method.

<!-- grade_id: cell-fe6400b434ccd5de, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-db55801f0b5063ae, points: 0, grade: true, solution: false -->
```python
#Do not make changes to this cell!
assert abbreviate("Adam Smith")=="A. Smith"
```

---
## Exercise 5: Palindromes

<b>Exercise 5A:</b> Create a function called `inv_string` that takes a string and displays it reversed.

For example, `inv_string("Python")` should return "nohtyP"

Hint: You can reverse a string using a slice with a negative step `[::-1]`.

<!-- grade_id: cell-2d10c118fe042823, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-969e1310a7aa8220, points: 1, grade: true, solution: false -->
```python
#Do not make changes to this cell!
```

---
<b>Exercise 5B:</b> A palindrome is a word, number, phrase, or other sequence of symbols that reads the same backwards as forwards.
Write a function named `is_palindrome` that checks if a word is a palindrome.
The following words are examples of palindromes: rotator, deified, racecar and reviver.

<!-- grade_id: cell-b70a5dc9d4d327f2, points: 0, grade: false, solution: true -->
```python
def is_palindrome(string):
    # YOUR CODE HERE
```

<!-- grade_id: cell-7c2e862d3f7d55c2, points: 1, grade: true, solution: false -->
```python
#Do not make changes to this cell!
assert is_palindrome("rotor")==True,"Incorrect"
assert is_palindrome("rotors")==False,"Incorrect"
assert is_palindrome("1234321")==True,"Incorrect"
assert is_palindrome("1234567")==False,"Incorrect"
```

---
<b>Exercise 5C:</b> Modify your `'is_palindrome'` function below to make sure it also accepts sentences. The function should correctly classify the following sentences as palindromes:</br>
- "A man, a plan, a canal – Panama"</br>
- "Doc, note: I dissent. A fast never prevents a fatness. I diet on cod"

<!-- grade_id: cell-1ea0e8e6317cc120, points: 0, grade: false, solution: true -->
```python
def is_palindrome(string):
    # YOUR CODE HERE
```

<!-- grade_id: cell-e22b076e7c8ae8c3, points: 1, grade: true, solution: false -->
```python
#Do not make changes to this cell!
assert is_palindrome("A man, a plan, a canal – Panama")==True,"Incorrect"
assert is_palindrome("Doc, note: I dissent. A fast never prevents a fatness. I diet on cod")==True,"Incorrect"
```

---
## Exercise 6

You have received a list of participants’ names for a workshop. As is often the case with user-entered data, the names may be inconsistently formatted. Your task is to write a function that properly formats the names.</br>
</br>
Examples:</br>
Femke Van Houten should become Femke van Houten</br>
ANNA Van Der Meer should become Anna van der Meer</br>
Franz h. BECKER should become Franz H. Becker</br>
Emily taylor johnson should become Emily Taylor Johnson</br>

The following prepositions should be printed in lower case letters only:
van, de, der, den and von.</br>
</br>
The function should be named `format_name`.

Hints:
- Start by creating an empty list named `styled` and a list named `prepositions` containing the prepositions above.
- Iterate over each word in the name: if it is in `prepositions`, add it to styled in lowercase; otherwise, add it in title case.
- You can combine the words in `styled` into a single string by creating a variable, for example: `formatted_name = ' '.join(styled)`




<!-- grade_id: format_names, points: 0, grade: false, solution: true -->
```python
def format_name(full_name):
    # YOUR CODE HERE
```

<!-- grade_id: format_names_test, points: 1, grade: true, solution: false -->
```python
#Do not make changes to this cell!
assert format_name('emily taylor Johnson') == 'Emily Taylor Johnson'
assert format_name('MICHAEL DAVIS') == 'Michael Davis'
assert format_name('Pieter De Vries') == 'Pieter de Vries'
assert format_name('lieke paula verbeek') == 'Lieke Paula Verbeek'
assert format_name('Lotte Van Den Berg') == 'Lotte van den Berg'
```

---
&copy; Jeroen van Raak, University of Amsterdam, 2025